# CC5215: Privacidad de Datos
## Laboratorio 3
**Integrantes**

- Nombre apellido

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
from unittest.mock import patch

plt.style.use('seaborn-v0_8-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

adult = pd.read_csv('https://users.dcc.uchile.cl/~mtoro/cursos/cc5215/adult_with_pii.csv')
adult = adult.dropna()
```

```python
adult.head()
```

## Question 1 (5 points)

Write code to answer the query: "how many participants have never been married?"

*Hint*: filter the `adult` dataframe to contain only participants who were never married, then return the  `len` of the filtered dataframe.

<!-- grade_id: cell-975eb03979d78eaf, points: 0, grade: false, solution: true -->
```python
def never_married_count_query():
    # YOUR CODE HERE
    raise NotImplementedError()

never_married_count_query()
```

<!-- grade_id: cell-c5ba2c93a46e8c5e, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1
assert never_married_count_query() == 9726
```

## Question 2 (5 points)

In 2-5 sentences, answer the following:
- What is the sensitivity of `never_married_count_query`, and why?

<!-- grade_id: cell-fb1c7b0533f933f7, points: 5, grade: true, solution: true -->

ANSWER:

## Question 3 (5 points)

Use the implementation of `laplace_mech` to produce a differentially private answer to `never_married_count_query`, with `epsilon = 0.1`.

<!-- grade_id: cell-80d3c108ba0f75d5, points: 0, grade: false, solution: true -->
```python
def dp_never_married_count_query(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

dp_never_married_count_query(0.1)
```

<!-- grade_id: cell-93eab43d27806309, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3
dp_results = [dp_never_married_count_query(0.1) for _ in range(100)]
spec = [np.random.laplace(loc=9726, scale=1/0.1) for _ in range(100)]
assert stats.wasserstein_distance(dp_results, spec) < 10
```

# Question 4 (10 points)

Write the implementation of a class `DPPrinter` that collects the results of differentially private computations along with the used epsilon. Also, implement the method `print` that prints the results collected along with the total epsilon.

```python
class DPPrinter():
    results = []
    epsilon = 0

    def push(self, result, eps):
        """
        Receives a differentially private result and the epsilon used to
        privatize the original result, and it stores them in the class
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def print(self):
        """
        Prints the results collected so far and the total epsilon
        """
        # YOUR CODE HERE
        raise NotImplementedError()

printer = DPPrinter()

printer.push(dp_never_married_count_query(0.1), 0.1)
printer.push(dp_never_married_count_query(1.0), 1.0)
printer.push(dp_never_married_count_query(0.1), 0.1)

printer.print()
```

```python
assert len(printer.results) == 3
assert np.isclose(printer.epsilon, 1.2)
```

## Question 5 (10 points)

The `pct_error` function, defined below, returns the percent relative error between an original query result and a differentially private result for the same query.

Implement a function `graph_error_never_married` that:

- Calculates 1000 differentially private answers to `dp_never_married_count_query`
- Calculates the percent error for each one of these answers against the original (non-private) answer
- Graphs the distribution of errors using a histogram

*Hint*: use `plt.hist(..., bins=20)`.

The given code will use your function to plot errors for `epsilon=0.1` and `epsilon=1.0`.

```python
def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0
```

<!-- grade_id: cell-da0c198a1cf9a866, points: 0, grade: false, solution: true -->
```python
def graph_error_never_married(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

graph_error_never_married(0.1)
graph_error_never_married(1.0)
```

<!-- grade_id: cell-eda6bc27840a9067, points: 10, grade: true, solution: false -->
```python
error_spec = 100.0 * np.abs(np.random.laplace(loc=0, scale=1, size=1000)) / never_married_count_query()

with patch('matplotlib.pyplot.hist') as mock_hist:
    graph_error_never_married(1.0)

args, kwargs = mock_hist.call_args
assert stats.wasserstein_distance(error_spec, args[0]) < 0.001
```

## Question 6 (10 points)

In 2-5 sentences, answer the following:

1. How does the histogram of relative errors for $\epsilon = 0.1$ differ from the one for $\epsilon = 1.0$?
2. What do the two histograms tell you about the effect of $\epsilon$ on relative error?

**HINT**: Pay close attention to which color represents which value of $\epsilon$!

<!-- grade_id: cell-075013c43cc4cc1d, points: 10, grade: true, solution: true -->
ANSWER 1:

ANSWER 2:

## Question 7 (5 points)

Consider `over_60_count_query`, which asks how many people in the dataset are over the age of 60.

```python
def over_60_count_query():
    return len(adult[adult['Age'] > 60])
```

Implement `dp_over_60_count_query`, a differentially private version of `over_60_count_query` (as in question 3), and `graph_error_over_60`, which graphs relative error for `dp_over_60_count_query` (as in question 5).

<!-- grade_id: cell-1c15ba7324d807a4, points: 0, grade: false, solution: true -->
```python
def dp_over_60_count_query(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

def graph_error_over_60(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

# we plot both errors for never_married_count_query and over_60_count_query
# at the same epsilon, to compare
graph_error_over_60(1.0)
graph_error_never_married(1.0)
```

<!-- grade_id: cell-239ae948aa08f924, points: 10, grade: true, solution: false -->
```python
error_spec = 100.0 * np.abs(np.random.laplace(loc=0, scale=1, size=1000)) / over_60_count_query()

with patch('matplotlib.pyplot.hist') as mock_hist:
    graph_error_over_60(1.0)

args, kwargs = mock_hist.call_args
assert stats.wasserstein_distance(error_spec, args[0]) < 0.01
```

## Question 8 (10 points)

In 2-5 sentences, answer the following:

1. Given the graph from question 5, how does relative error differ between `dp_never_married_count_query` and `dp_over_60_count_query` for the same value of $\epsilon$?
2. What property of the query causes the difference in relative errors between `dp_never_married_count_query` and `dp_over_60_count_query`?

<!-- grade_id: cell-be5745d92a22154c, points: 10, grade: true, solution: true -->
ANSWER 1:

ANSWER 2: